Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Hung-Wei Chang"
COLLABORATORS = ""
```

---

<a href='https://ai.meng.duke.edu'> = <img align="left" style="padding-top:10px;" src=https://storage.googleapis.com/aipi_datasets/Duke-AIPI-Logo.png>

# Assignment 8

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import statsmodels.api as sm
from statsmodels.formula.api import ols
from scipy.stats import ttest_ind
import seaborn as sns

import warnings
warnings.filterwarnings("ignore")
```

## Question 1
In this question we will be working with three sets of data related to patients who have been diagnosed with severe tumors.  Some patients have received a new experimental drug and others have not.  Our objective is to analyze the effect that the drug has had on mean survival time of the patients.  

We have the following columns in our data (after running the below cell to rename):  
- `Patient`: patient id number. 
- `Received`: whether the patient received the new drug (1) or did not (2). 
- `Diagnosis`: whether the initial tumor diagnosis was a state 3 (1) or 4 (2). 
- `Days`: the survival time in days from time of diagnosis

```python
# Read in and format the data
data_mar = pd.read_csv('Data_MAR.csv')
data_mcar = pd.read_csv('Data_MCAR.csv')
data_mnar = pd.read_csv('Data_MNAR.csv')

for df in [data_mar,data_mcar,data_mnar]:
    df.rename({'Patient ID Number': 'Patient',
                          'Received New Medicine Yes  (1)   No (2)':'Received',
                          'Initial Tumor Diagnosis State 3-4':'Diagnosis',
                         'Survival Time - Days': 'Days'}, axis=1, inplace=True)
    
data_mar.head()
```

Output:
```
Patient  Received  Diagnosis  Other major illness   Days
0        1         2          1                  4.0    NaN
1        2         1          1                  5.0    NaN
2        3         2          2                  5.0  789.0
3        4         1          2                  5.0  700.0
4        5         2          1                  5.0    NaN
```

### Q1.1
Now, complete the below function `sample_stats_control()` which takes as input a dataframe of patient data and calculates basic statistics for the control group (the patients who did not receive the new drug).  The function should first remove any rows with null values in the `Days` column.  It should then calculate and return the sample size for the control group (number of patients who did not receive the medicine) as an integer and the mean survival time (`Days`) for the control group as a float.

<!-- grade_id: cell-e321e38ab586ed00, points: 0, grade: false, solution: true -->
```python
def sample_stats_control(df):
    # YOUR CODE HERE
    df_test = df.copy()
    df_final = df_test.loc[ (df_test['Received'] == 2 ) & (  ~df_test['Days'].isnull() )]
    samp_size = int(len(df_final.index))
    mean_surv = float(df_final['Days'].mean())
    return samp_size, mean_surv
    raise NotImplementedError()
```

<!-- grade_id: cell-243d12f7298bc8b5, points: 16, grade: true, solution: false -->
```python
# Test cell for sample_stats_control()
for df in [data_mar,data_mcar,data_mnar]:
    size,mean_days = sample_stats_control(df)
    print('For the control group in dataset {} your function calculated the following:'.
          format([x for x in globals() if globals()[x] is df][0]))
    print('Sample size: {}'.format(size))
    print('Sample mean: {:.2f}'.format(mean_days))
    
```

Output:
```
For the control group in dataset data_mar your function calculated the following:
Sample size: 91
Sample mean: 507.65
For the control group in dataset data_mcar your function calculated the following:
Sample size: 72
Sample mean: 556.56
For the control group in dataset data_mnar your function calculated the following:
Sample size: 90
Sample mean: 503.71
```

### Q1.2
Now, complete the below function `sample_stats_treatment()` which takes as input a dataframe of patient data and calculates basic statistics for the group of patients who did receive the new drug.  The function should first remove any rows with null values in the `Days` column.  It should then calculate and return the sample size for the group of patients who received the drug as an integer, and the mean survival time (`Days`) for this group as a float.

<!-- grade_id: cell-af55d18cd887907a, points: 0, grade: false, solution: true -->
```python
def sample_stats_treatment(df):
    # YOUR CODE HERE
    df_test = df.copy()
    df_final = df_test.loc[ (df_test['Received'] == 1 ) & (  ~df_test['Days'].isnull() )]
    samp_size = int(len(df_final.index))
    mean_surv = float(df_final['Days'].mean())
    return samp_size, mean_surv

    raise NotImplementedError()
```

<!-- grade_id: cell-7ade43d5be18a592, points: 16, grade: true, solution: false -->
```python
# Test cell for sample_stats_treatment()
for df in [data_mar,data_mcar,data_mnar]:
    size,mean_days = sample_stats_treatment(df)
    print('For the treatment group in dataset {} your function calculated the following:'.
          format([x for x in globals() if globals()[x] is df][0]))
    print('Sample size: {}'.format(size))
    print('Sample mean: {:.2f}'.format(mean_days))
    
```

Output:
```
For the treatment group in dataset data_mar your function calculated the following:
Sample size: 89
Sample mean: 592.53
For the treatment group in dataset data_mcar your function calculated the following:
Sample size: 63
Sample mean: 665.87
For the treatment group in dataset data_mnar your function calculated the following:
Sample size: 100
Sample mean: 648.48
```

### Q1.3
Let's now calculate the p-value of the new medicine against the control group for each dataset.  Complete the below function `calc_pval()` which takes a dataframe of patient data as input and determines if there is a statistically significant difference in mean survival time (`Days`) of the patients receiving the medicine versus those that do not receive it, using a two-sided Student's t-test, at an alpha=0.05.  Be sure to first remove any rows with null values for `Days`.  Your function should return the p-value of the t-test as a float.

<!-- grade_id: cell-df4ea3bc0c109c29, points: 0, grade: false, solution: true -->
```python
def calc_pval(df):
    # YOUR CODE HERE
    df_test = df.copy()
    col_1 = df_test.loc[ (df_test['Received'] == 2 ) & (  ~df_test['Days'].isnull() )]['Days']
    col_2 = df_test.loc[ (df_test['Received'] == 1 ) & (  ~df_test['Days'].isnull() )]['Days']
    results = ttest_ind(col_1, col_2)
    return float(results.pvalue)
    
    raise NotImplementedError()
```

<!-- grade_id: cell-bb22cbc85ddecbb4, points: 20, grade: true, solution: false -->
```python
# Test cell for calc_pval()
for df in [data_mar,data_mcar,data_mnar]:
    p_value = calc_pval(df)
    print('For the dataset {} your function calculated the p-value as {:.4f}'.
          format([x for x in globals() if globals()[x] is df][0],p_value))
    
```

Output:
```
For the dataset data_mar your function calculated the p-value as 0.0142
For the dataset data_mcar your function calculated the p-value as 0.0205
For the dataset data_mnar your function calculated the p-value as 0.0001
```

## Question 2
Each of the three data sets A, B, C are drawn from one of the three continuous distributions discussed in class: gaussian, exponential and continuous distributions.  Our objective in this question is to identify the underlying distribution behind each of these datasets and the maximum likelihood parameters(s) for each.


```python
# Read in data and separate it into the 3 datasets
data = pd.read_csv('data_kde.csv')
data_A = data['Data Set A']
data_B = data['Data Set B']
data_C = data['Data Set C']
```

### Q2.1
Complete the below function `analyze_A()` which takes as input the `data_A` dataset.  Your function should visualize the data using a histogram and/or kde plot to evaluate the distribution.  Your function should then return the following:  
- The name of the distribution as a string (choose from 'gaussian','exponential' or 'continuous'. 
- A **list** containing the key parameter or parameters required to fully describe the distribution

<!-- grade_id: cell-b9342fe3183b934f, points: 0, grade: false, solution: true -->
```python
def analyze_A(df):
    # YOUR CODE HERE
    param_list = [ float ( 1 / df.mean() ) ] 
    return 'exponential', param_list
    raise NotImplementedError()
```

```python
# df = data_C
# # View histogram of each feature amygdala and acc
# fig,ax = plt.subplots(2,2,figsize=(15,10))
# ax[0,0].hist(df, bins=5)
# ax[0,0].set_title('amygdala, bins=5')
# ax[0,1].hist(df,bins=5)
# ax[0,1].set_title('acc, bins=5')
# ax[1,0].hist(df,bins=20)
# ax[1,0].set_title('amygdala, bins=20')
# ax[1,1].hist(df,bins=20)
# ax[1,1].set_title('acc, bins=20')
# plt.show()
```

```python
# # View KDE plot of amygdala and acc
# df = [data_A, data_B, data_C]
# for i in df:
#     sns.kdeplot(i)
#     plt.show()
```

<!-- grade_id: cell-63fb29cc9496d2c7, points: 16, grade: true, solution: false -->
```python
# Test cell for analyze_A
dist,params = analyze_A(data_A)
print('Your function identified the distribution as: {}'.format(dist))
print('Your function identified the parameter(s) as: {}'.format([np.round(i,2) for i in params]))
assert len(params)==1
assert type(dist) == str

```

Output:
```
Your function identified the distribution as: exponential
Your function identified the parameter(s) as: [0.09]
```

### Q2.2
Complete the below function `analyze_B()` which takes as input the `data_B` dataset.  Your function should visualize the data using a histogram and/or kde plot to evaluate the distribution.  Your function should then return the following:  
- The name of the distribution as a string (choose from 'gaussian','exponential' or 'continuous'. 
- A **list** containing the key parameter or parameters required to fully describe the distribution (in any order)

<!-- grade_id: cell-c4ab99ea3925f1ef, points: 0, grade: false, solution: true -->
```python
def analyze_B(df):
    # YOUR CODE HERE
    mean = df.mean()
    stddev = df.std()
    param_list = [mean, stddev]
    return 'gaussian', param_list
    
    raise NotImplementedError()
```

<!-- grade_id: cell-2bc084275bcc4553, points: 16, grade: true, solution: false -->
```python
# Test cell for analyze_B
dist,params = analyze_B(data_B)
print('Your function identified the distribution as: {}'.format(dist))
print('Your function identified the parameter(s) as: {}'.format([np.round(i,2) for i in params]))
assert len(params)==2
assert type(dist) == str

```

Output:
```
Your function identified the distribution as: gaussian
Your function identified the parameter(s) as: [14.08, 3.67]
```

### Q2.3
Complete the below function `analyze_C()` which takes as input the `data_C` dataset.  Your function should visualize the data using a histogram and/or kde plot to evaluate the distribution.  Your function should then return the following:  
- The name of the distribution as a string (choose from 'gaussian','exponential' or 'continuous'. 
- A **list** containing the key parameter or parameters required to fully describe the distribution (in any order)

<!-- grade_id: cell-78a1af671fcb478e, points: 0, grade: false, solution: true -->
```python
def analyze_C(df):
    # YOUR CODE HERE
    params_list = [df.min(), df.max()  ]
    return 'continuous', params_list

    raise NotImplementedError()
```

<!-- grade_id: cell-a1ec8b1bc371987c, points: 16, grade: true, solution: false -->
```python
# Test cell for analyze_C
dist,params = analyze_C(data_C)
print('Your function identified the distribution as: {}'.format(dist))
print('Your function identified the parameter(s) as: {}'.format([np.round(i,2) for i in params]))
assert len(params)==2
assert type(dist) == str

```

Output:
```
Your function identified the distribution as: continuous
Your function identified the parameter(s) as: [-48.84, 49.25]
```

```python

```